## SLU08 - Functions Intermediate: Exercise notebook

```python
# Make sure to run this cell or the grading won't work!

import math
import hashlib
import inspect
import numpy as np
from random import randint
from utils import check_call
```

### Exercise 1 (ungraded)

Open your favorite text editor and write an answer to the following questions in your own words:

1. Explain what are scopes.
1. Explain what are namespaces.
1. Explain what is local scope.
1. Explain what is global scope.

YOUR ANSWER HERE

### Exercise 2
Consider the code in the cell below and identify the scopes of the variables.

```python
def hello_world():
    a_var = 10
    return a_var

another_var = 100
```

<!-- grade_id: cell-954725562970a8a3, points: 0, grade: false, solution: true -->
```python
# assign the string values of 'local' or 'global' to each answer
# a_var_scope = 
# another_var_scope = 

### BEGIN SOLUTION
a_var_scope = 'local'
another_var_scope = 'global'
### END SOLUTION`
```

<!-- grade_id: cell-5beb29cb41412c43, points: 2, grade: true, solution: false -->
```python
assert hashlib.sha256(a_var_scope.encode()).hexdigest() == '25bf8e1a2393f1108d37029b3df5593236c755742ec93465bbafa9b290bddcf6', 'Incorrect answer for a_var_scope, try again!'
assert hashlib.sha256(another_var_scope.encode()).hexdigest() == '8001c27439650c5c5a6b4ed94163b5ddeb4476362c71380e613fa20dfffcef50', 'Incorrect answer for another_var_scope, try again!'
print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

### Exercise 3
Consider the code below and find out what will be printed by the four `print` statements. Think about it without running the code.

```python
def parallel_universe():
    a = 10
    print(a) # first print

    def hello_universe():
        print(a) # second print

    def good_morning_sunshine(a):
        print(a) # third print
        a = 50
    
    hello_universe()
    good_morning_sunshine(a)

a = 200
parallel_universe()
print(a)  # fourth print
```

<!-- grade_id: cell-f108f9b2b1624025, points: 0, grade: false, solution: true -->
```python
# Assign the output of the corresponding print statement to these variables:
# print_1 = 
# print_2 = 
# print_3 = 
# print_4 = 

### BEGIN SOLUTION
print_1 = 10
print_2 = 10
print_3 = 10
print_4 = 200

### END SOLUTION
```

<!-- grade_id: cell-6f0bede927d6d49f, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(str(print_1).encode()).hexdigest() == '4a44dc15364204a80fe80e9039455cc1608281820fe2b24f1e5233ade6af1dd5', 'print_1 is not correct.'
assert hashlib.sha256(str(print_2*2).encode()).hexdigest() == 'f5ca38f748a1d6eaf726b8a42fb575c3c71f1864a8143301782de13da2d9202b', 'print_2 is not correct.'
assert hashlib.sha256(str(print_3*3).encode()).hexdigest() == '624b60c58c9d8bfb6ff1886c2fd605d2adeb6ea4da576068201b6c6958ce93f4', 'print_3 is not correct.'
assert hashlib.sha256(str(print_4*4).encode()).hexdigest() == '1a1cf797fabe7f95836fabeca626907c77b3e6c9aff7c2290b396a238c69362e', 'print_4 is not correct.'
print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

### Exercise 4 (ungraded)
Consider the code below. This code does not work. Think about why and write down you answer in the cell below.

```python
outside_var = 100

def hello_world_complicated():
    print(outside_var) # first print
    a_var = 10
    outside_var = 200
    print(outside_var) # second print

hello_world_complicated()
```

YOUR ANSWER HERE

### Exercise 5

Now you'll see first hand how local and global scope can clash (and potentially create problems).

Consider the code below. Without running the code, think about what will be the values of the variables `my_value_1` and `my_value_2`. Fill in the values of both lists in your answer.

```python
a_list = [1, 2, 3]

def hello():
    another_list = a_list
    another_list.append(4)
    
my_value_1 = a_list + [5]
hello()
my_value_2 = a_list
```

<!-- grade_id: cell-eb4e9cfc837086c7, points: 0, grade: false, solution: true -->
```python
# fill in the values of the list variables 
my_value_1 = [] 
my_value_2 = []

### BEGIN SOLUTION
my_value_1 = [1, 2, 3, 5]
my_value_2 = [1, 2, 3, 4]
### END SOLUTION
```

<!-- grade_id: cell-6938960a9a0f275f, points: 4, grade: true, solution: false -->
```python
assert isinstance(my_value_1,list)
assert hashlib.sha256(str(len(my_value_1)).encode()).hexdigest() == '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a', 'The length of the list my_value_1 is not correct.'
assert hashlib.sha256(str(min(my_value_1)).encode()).hexdigest() == '6b86b273ff34fce19d6b804eff5a3f5747ada4eaa22f1d49c01e52ddb7875b4b', 'The values in the list my_value_1 are not correct.'
assert hashlib.sha256(str(max(my_value_1)).encode()).hexdigest() == 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d', 'The values in the list my_value_1 are not correct.'
assert hashlib.sha256(str(sum(my_value_1)).encode()).hexdigest() == '4fc82b26aecb47d2868c4efbe3581732a3e7cbcc6c2efb32062c08170a05eeb8', 'The values in the list my_value_1 are not correct.'
assert hashlib.sha256(str(np.mean(my_value_1)).encode()).hexdigest() == '73af98c887a3fd39bb0bf46ca51699e9c8c250db05a7452ee522614ff657b886', 'The values in the list my_value_1 are not correct.'
assert isinstance(my_value_2,list)
assert hashlib.sha256(str(len(my_value_2)).encode()).hexdigest() == '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a', 'The length of the list my_value_2 is not correct.'
assert hashlib.sha256(str(min(my_value_2)).encode()).hexdigest() == '6b86b273ff34fce19d6b804eff5a3f5747ada4eaa22f1d49c01e52ddb7875b4b', 'The values in the list my_value_2 are not correct.'
assert hashlib.sha256(str(max(my_value_2)).encode()).hexdigest() == '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a', 'The values in the list my_value_2 are not correct.'
assert hashlib.sha256(str(sum(my_value_2)).encode()).hexdigest() == '4a44dc15364204a80fe80e9039455cc1608281820fe2b24f1e5233ade6af1dd5', 'The values in the list my_value_2 are not correct.'
assert hashlib.sha256(str(np.mean(my_value_2)).encode()).hexdigest() == 'b8736b999909049671d0ea075a42b308a5fbe2df1854899123fe09eb0ee9de61', 'The values in the list my_value_2 are not correct.'

print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

### Exercise 6.1
Create a function with the following specifications:
- Create a function called `use_global_scope` that takes two arguments
- The first argument is called `multiply_by` and should be an integer
- The second argument is called `a_number` and is a number
- The function should return the value of `a_number` multiplied by `multiply_by`
- If `multiply_by` is not an integer, the function returns `None`

<!-- grade_id: cell-677ed0e5bc6a02ac, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def use_global_scope(multiply_by, a_number):
    if type(multiply_by) is not int:
        return None
    
    return multiply_by * a_number
### END SOLUTION
```

<!-- grade_id: cell-eb4e97d47bbb5aa8, points: 4, grade: true, solution: false -->
```python
# check function source
sig = inspect.signature(use_global_scope)
source = inspect.getsource(use_global_scope)
source_lines = source.split('\n')
params = sig.parameters

# check function name
assert ' '.join(source_lines[0].split('(')[0].split())=='def use_global_scope', 'The function name is not correct.' 

# check function parameters
assert len(params) == 2, 'The function should take two parameters.'
assert list(params.keys())[0] == 'multiply_by', 'The name of the first parameter is not correct.'
assert list(params.keys())[1] == 'a_number', 'The name of the second parameter is not correct.'

# test function
i1=randint(0,100)
i2=randint(0,100)
assert use_global_scope(i1,i2) == i1*i2, 'The function is not multiplying correctly.'
i1=randint(0,100)
i2=randint(0,100)/4
assert use_global_scope(i1,i2) == i1*i2, 'The function is not multiplying correctly.'
assert use_global_scope(1.45,6739) == None, 'If multiply_by is not an integer, the function should return None.'

print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

### Exercise 6.2 
Now use the function you've defined in 6.1:
- Create a global variable called `a_number` and assign it the value of `2`.
- Call the `use_global_scope` function, passing one positional and one keyword argument.
- Use the global variable as the value of the `a_number` argument.
- Use 14 as the value of the `multiply_by` argument.

<!-- grade_id: cell-d77f348198eb2cd6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a_number = 2
use_global_scope(14, a_number = a_number)
### END SOLUTION
```

Output:
```
28
```

<!-- grade_id: cell-9ea37f9eaa3577b0, points: 3, grade: true, solution: false -->
```python
# check the a_number variable
'a_number' in globals().keys(), 'The variable a_number does not exist.'
assert a_number == 2, 'The value of a_number should be 2.'

# check the function call
keys_list=list(globals().keys())
l_keys=len(keys_list)
output=' '.join([str(globals()[keys_list[i]]) for i in range(l_keys-10,l_keys)])
output_w_spaces=''.join(output.split())
assert check_call(output_w_spaces), 'Are you calling the function as specified?'

print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

### Exercise 7

In this exercise, you will create a function that multiplies all the elements of a list
by a given number.

- Create a function named `multiply_list`.
- The function should have two parameters, `my_list` and `multiply_by`.
- The `my_list` parameter is a list.
- The `multiply_by` parameter is a number and its default value is `2`.
- The function should multiply all the elements of `my_list` by `multiply_by` and return the result (which will also be a list).
- The function should not modify the input list.

<!-- grade_id: cell-01eeb359131095df, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def multiply_list(my_list, multiply_by=2):
    multiplied_by = []
    for num in my_list:
        multiplied_by.append(num * multiply_by)
    return multiplied_by
### END SOLUTION
```

<!-- grade_id: cell-1816949c415ea5af, points: 4, grade: true, solution: false -->
```python
# check function source
sig = inspect.signature(multiply_list)
source = inspect.getsource(multiply_list)
params = sig.parameters
source_lines = source.split('\n')

# check function name
assert ' '.join(source_lines[0].split('(')[0].split())=='def multiply_list', 'The function name is not correct.' 

# check function parameters
assert len(params) == 2, 'The function should take two parameters.'
assert 'my_list' in params, 'The function should a parameter called my_list.'
assert 'multiply_by' in params, 'The function should have a parameter called multiply_by.'
assert params['multiply_by'].default == 2, 'The multiply_by parameter should have a default value of 2.'

# test function
test_list=[randint(0,100) for i in range(randint(1,10))]
sum_list=sum(test_list)
result_list = multiply_list(my_list=test_list,multiply_by=13) 
assert sum_list == sum(test_list), 'The function should not modify the input list.'
test_list=[randint(0,100) for i in range(randint(1,10))]
result_list = multiply_list(test_list)
assert sum(result_list) == sum(test_list)*2, 'The function is not multiplying correctly.'
test_list=[randint(0,100) for i in range(randint(1,10))]
result_list = multiply_list(my_list=test_list,multiply_by=12) 
assert sum(result_list) == sum(test_list)*12, 'The function is not multiplying correctly.'

print('Very good, your answer is correct!')
```

Output:
```
Very good, your answer is correct!
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)