# SLU18 - Exam preparation II

This is the wave 1 Python admission exam from the Lisbon Data Science Starters Academy batch 6. The allocated time for this exam was 3 hours. If you plan to take the admission exam this year, it's a good idea to measure the time you need to solve this notebook. You can see how far you get in 3 hours and also how much time you need to complete the notebook. At the same time, it's a good idea to not rush through the exercises and rather concentrate on working thoroughly.

At the end, you submit the notebook to the portal as usual. You can submit as many times as you like.

```python
import csv
import numpy as np
import inspect
```

## Exercise 1 (5 points)

This exercise is a quiz with 10 multiple choice questions.
In the cell below, we've declared a dictionary called `answers`.

You should fill in that dictionary with your answers, using as keys the question numbers, like `question_x`, and as values the numbers from `1` to `10` that correspond to the right answer.
For each question, you should provide only one answer, *i.e.*, the dict values should have type **int**.

For example, if you want to answer Question 1 with choice number 2, then you do:
```
answers["question_1"] = 2
```

```python
answers = {}
```

### Part I (questions 1-5)

Imagine that you are using a Unix based machine, and that you have a terminal with the working directory `/users/john`. This directory has in it:
- a file named `example_data.csv`;
- two directories, `documents` and `ldsa`, containing some files.

The file tree below illustrates the contents and structure of `/users/john`, **that should be considered to answer Questions 1-5**.


```text
/users/john
├── example_data.csv
├── documents
│   ├── awesome_profile_picture.png
│   ├── cv.pdf
│   └── todo.txt
└── ldsa
    ├── example.py
    ├── notes.txt
    └── .gitignore
```

#### Question 1
How would you list the contents of the current working directory?
1. `pwd`
2. `cd .`
3. `ls documents`
4. `ls .`

<!-- grade_id: cell-98b6eeed9f5eb544, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_1"] = 4
### END SOLUTION

# answers["question_1"] = ...
```

#### Question 2
Which command would you use to see the entire content of `example_data.csv`?
1. `head -5 example_data.csv`
2. `cat example_data.csv`
3. `head -5 documents/example_data.csv`
4. `cat documents/example_data.csv`

<!-- grade_id: cell-f75169d418f280ba, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_2"] = 2
### END SOLUTION

# answers["question_2"] = ...
```

#### Question 3
After running one of the following commands, there will be 4 files in the `ldsa` directory. Which command is it?
1. `cp example_data.csv ldsa`
2. `rm -r example_data.csv ldsa`
3. `touch example_data.csv`
4. `cp ldsa example_data.csv`

<!-- grade_id: cell-d0ee1b32bb06368b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_3"] = 1
### END SOLUTION

# answers["question_3"] = ...
```

#### Question 4
The command `rm documents/awesome.png`
1. Creates a new file in the `documents` directory
2. Removes a file from the `documents` directory
3. Removes the entire `documents` directory
4. Raises an error

<!-- grade_id: cell-b48ab5ac85e1bf69, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_4"] = 4
### END SOLUTION

# answers["question_4"] = ...
```

#### Question 5
Based on the structure of `/users/john` and its subdirectories, which folder is likely to be using git to version control its contents? 
1. `/users/john`
2. `/users/john/documents`
3. `/users/john/ldsa`
4. None of the above

<!-- grade_id: cell-d53190f9ce4dedbe, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_5"] = 3
### END SOLUTION

# answers["question_5"] = ...
```

### Part II (questions 6-10)

Answer the following questions about git and Python programming.

#### Question 6
Consider the following git commands: `add`, `commit`, `status`, `pull`, `push`. Which is the most natural order to use them?
1. `pull` > `status` > `add` > `commit` > `push`
2. `status` > `commit` > `add` > `push` > `pull`
3. `push` > `add` > `status` > `commit` > `pull`
4. `pull` > `push` > `status` > `add` > `commit`

<!-- grade_id: cell-44e7a38840f54f6c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_6"] = 1
### END SOLUTION

# answers["question_6"] = ...
```

#### Question 7
Which one of the following sentences is true?

1. Dictionaries can have duplicated keys.
2. Dictionaries can have lists as keys.
3. Lists can have duplicated elements.
4. Tuple elements can be updated.

<!-- grade_id: cell-8f6345404f52ff67, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_7"] = 3
### END SOLUTION

# answers["question_7"] = ...
```

#### Question 8
Consider that `people` is a list of instances of a class `Person`. Which of the following instructions builds a list of the attributes `x` of those instances? 
1. `[person["x"] for person in people]`
2. `(person["x"] for person in people)`
3. `[person.x for person in people]`
4. `(person.x for person in people)`

<!-- grade_id: cell-67e3948e70b7e795, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_8"] = 3
### END SOLUTION

# answers["question_8"] = ...
```

#### Question 9
A module `my_module` has a function `my_function`. Which sequence of commands allows you to use `my_function` from a file in the same directory as `my_module.py`?

1. `import my_function` and `y = my_module()`
2. `import my_module` and `y = my_function()`
3. `from my_function import my_module` and `y = my_module()`
4. `from my_module import my_function` and `y = my_function()`

<!-- grade_id: cell-fd120313a3617c0f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_9"] = 4
### END SOLUTION

# answers["question_9"] = ...
```

#### Question 10
Consider that the variable `a` stores the float `2.0`. After which of the following operations does the assert `a == 4.` fail?
1. `a = a * a`
2. `a += 2.`
3. `a *= 2.`
4. `a /= 2.`

<!-- grade_id: cell-03d49ee095b85596, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_10"] = 4
### END SOLUTION

# answers["question_10"] = ...
```

<!-- grade_id: cell-1daf63702b1b0228, points: 5, grade: true, solution: false -->
```python
from hidden_tests import test_exercise_1
test_exercise_1(answers)
```

Output:
```
10 / 10 correct answers. Your score is 5.
5
```

## Exercise 2

Consider a csv file that stores information about popular movies. This file has four columns: `imdb_id`, `title`, `year`, `rating`.

File `movies.csv` is an example of such a file. You can preview it using some command line instructions.

### Part I (2 points)

Implement a function that reads a file with the same format as the `movies.csv` file, and stores the data in a dictionary with the following structure:

```
{
    imdb_id: {
        "title": movie title (type: str),
        "year": release year of the movie (type: int),
        "rating": number between 0 and 10 representing the current movie rating (type: float)
    },
}

```

The function should:
1. be called `read_movies`;
2. receive an argument called `file_path`, the path to the file that the function should read the data from;
3. return the dictionary that was created.

You should use Python's [`csv`](https://docs.python.org/3/library/csv.html) module to read the csv file. The module was already imported at the top of the Notebook.

<!-- grade_id: cell-622c6adc26159244, points: 0, grade: false, solution: true -->
```python
def read_movies(file_path):
    """
    Reads the file in file_path, parses it and returns the data in a dictionary.

    Parameters:
    file_path (str): Path to the input file to be parsed.

    Returns:
    movies_parsed_data (dict): Movies dataset stored as a dictionary, where the keys are IMDb IDs
        and the values are data about the movie.
    """
    
    ### BEGIN SOLUTION
    with open(file_path, "r") as f:
        reader = csv.reader(f, delimiter=",", quotechar="\"")

        next(reader)

        movies_parsed_data = {}    
        for i, line in enumerate(reader):
            movies_parsed_data[line[0]] = {
                "title": line[1],
                "year": int(line[2]),
                "rating": float(line[3])
            }
    
    return movies_parsed_data
    ### END SOLUTION
```

<!-- grade_id: cell-6c24a4ad6f30ce55, points: 2, grade: true, solution: false -->
```python
movies_data = read_movies("movies.csv")

source = inspect.getsource(read_movies)
assert 'csv.reader(' in source, 'You should use the csv module to read the file.'

assert isinstance(movies_data, dict)
assert len(movies_data) == 20

assert "tt0082971" in movies_data
indiana_jones = movies_data["tt0082971"]
assert indiana_jones["title"] == "Indiana Jones and the Raiders of the Lost Ark"
assert indiana_jones["year"] == 1981
np.testing.assert_almost_equal(indiana_jones["rating"], 8.4, decimal=1)


from more_tests import test_exercise_2_part_I
test_exercise_2_part_I(read_movies)

print('All asserts passed.')
```

Output:
```
All asserts passed.
```

### Part II (2 points)

Consider now that we want to use our dataset to compute the average rating of all movies released in a given year. 

Implement a function called `average_movie_rating` that:
1. receives as input
   - a dictionary, like the dictionary returned in Part I of this exercise; 
   - an integer representing the release year;
2. returns the average rating of all movies released in a given year.

If the release year provided as input does not exist in the dataset, the function should return `-1.`.

<!-- grade_id: cell-7854488ee93d67b6, points: 0, grade: false, solution: true -->
```python
def average_movie_rating(movies_data, year):
    """
    Computes the average rating of all movies released in a given year.

    Parameters:
    movies_data (dict): Dictionary with the movies dataset.
    year (int): Release year.

    Returns:
    average_movie_rating (float): Average rating of all movies released in a given year.
    """
    
    ### BEGIN SOLUTION
    from statistics import mean
    
    if year not in [movie["year"] for movie in movies_data.values()]:
        return -1.
    
    year_movie_ratings = [movie["rating"] for movie in movies_data.values() if movie["year"] == year]
    average_movie_rating = mean(year_movie_ratings)

    return average_movie_rating
    ### END SOLUTION
```

<!-- grade_id: cell-56eb4a74f3544d03, points: 2, grade: true, solution: false -->
```python
movies_data = {
    "tt0087843": {"title": "Once Upon a Time in America", "year": 1984, "rating": 8.3},
    "tt0066921": {"title": "A Clockwork Orange", "year": 1971, "rating": 8.3},
    "tt0435761": {"title": "Toy Story 3", "year": 2010, "rating": 8.3},
    "tt0084302": {"title": "The Marathon Family", "year": 1982, "rating": 8.9},
    "tt0080684": {"title": "Star Wars: Episode V - The Empire Strikes Back", "year": 1980, "rating": 8.7},
    "tt0086879": {"title": "Amadeus", "year": 1984, "rating": 8.4},
}


average_movie_rating_1980 = average_movie_rating(movies_data, 1980)
assert isinstance(average_movie_rating_1980, float)  
np.testing.assert_almost_equal(average_movie_rating_1980, 8.7, decimal=2)

average_movie_rating_1984 = average_movie_rating(movies_data, 1984)
assert isinstance(average_movie_rating_1984, float)  
np.testing.assert_almost_equal(average_movie_rating_1984, 8.35, decimal=2)

average_movie_rating_1992 = average_movie_rating(movies_data, 1992)
assert isinstance(average_movie_rating_1992, float)  
np.testing.assert_almost_equal(average_movie_rating_1992, -1., decimal=2)

from more_tests import test_exercise_2_part_II
test_exercise_2_part_II(average_movie_rating)

print('All asserts passed.')
```

Output:
```
All asserts passed.
```

### Part III (2 points)

Consider now that we want to investigate whether old movies tend to be better rated than recent ones.

To do that, we shall compute [Pearson's correlation coefficient](https://en.wikipedia.org/wiki/Pearson_correlation_coefficient) between the movies' release year and rating. This coefficient measures the strength of the linear relationship between two sets of data, *i.e.*, their linear correlation.

Create a function called `pearson_correlation_year_rating` that computes Pearson's correlation coefficient between the year and rating of all movies in our dataset. You may (and are strongly encouraged to) use the [`pearsonr`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.pearsonr.html) function from the Python package [`scipy`](https://docs.scipy.org/doc/scipy).

This function should:
1. receive as input a dictionary like the one we created in the Part I of this exercise;
2. return the Pearson's correlation coefficient between the year and rating of all movies in our dataset.

<!-- grade_id: cell-909dc68555453eea, points: 0, grade: false, solution: true -->
```python
def pearson_correlation_year_rating(movies_data):
    """
    Computes Pearson's correlation coefficient between movies' release years and ratings.

    Parameters:
    movies_data (dict): Dictionary with the movies dataset.

    Returns:
    pearson_correlation (float): Pearson's correlation coefficient.
    """
    
    ### BEGIN SOLUTION
    from scipy.stats import pearsonr
    
    movies_years = [movie["year"] for movie in movies_data.values()]
    movies_ratings = [movie["rating"] for movie in movies_data.values()]
    
    pearson_correlation, _ = pearsonr(movies_years, movies_ratings)

    return pearson_correlation
    ### END SOLUTION
```

<!-- grade_id: cell-8b56f5a17ceb13b6, points: 2, grade: true, solution: false -->
```python
movies_data = {
    "tt0087843": {"title": "Once Upon a Time in America", "year": 1984, "rating": 8.3},
    "tt0066921": {"title": "A Clockwork Orange", "year": 1971, "rating": 8.3},
    "tt0435761": {"title": "Toy Story 3", "year": 2010, "rating": 8.3},
    "tt0084302": {"title": "The Marathon Family", "year": 1982, "rating": 8.9},
    "tt0080684": {"title": "Star Wars: Episode V - The Empire Strikes Back", "year": 1980, "rating": 8.7},
    "tt0086879": {"title": "Amadeus", "year": 1984, "rating": 8.4},
}

pearson_correlation = pearson_correlation_year_rating(movies_data)
assert isinstance(pearson_correlation, float)
np.testing.assert_almost_equal(pearson_correlation, -0.243, decimal=3)

from more_tests import test_exercise_2_part_III
test_exercise_2_part_III(pearson_correlation_year_rating)

print('All asserts passed.')
```

Output:
```
All asserts passed.
```

## Exercise 3 (5 points)

In this exercise, we'll use object oriented programming concepts to model a human train, a dance choreography typical of Portuguese popular music &mdash; and one very much practiced in the recent [Santo António festivities in Lisbon](https://www.culturanarua.pt/en/program/lisbon-festivities/)! 🥳

In this choreography, there is one person leading the train, deciding its direction at all times. The second person on the train follows immediately behind the leader, the third person follows behind the second person, and so forth. During the choreography, people can join the train by following the last person in it.

The choreography is such that each person on the train immitates the trajectory of the person ahead of them, which is ultimately the trajectory of the train leader as well.

See [this video](https://www.youtube.com/watch?v=C1pB9pUQ3pQ) for an example of this choreography.

### Human train model

We shall assume that the train moves in a square grid. It can take one step at a time &mdash; forward, left or right. It cannot collide with itself.


When viewed from overhead, this movement can be described in a grid, of axes `x` and `y`. The position of a person on the train with label `i` is thus given by two coordinates, `[x_i, y_i]`, where `x_i` and `y_i` are integer numbers.

<img src="images/train_definition.png" alt="Human train definition" width="350"/>


We shall model this train as a sequence of people of which we track the positions `[x_i, y_i]`. An example of this model is illustrated in the figure above. The train leader is the first person in this sequence (`i=0`, highlighted in red) and is followed by all other people on the train `i=1, ..., N - 1`, where `N` is the total number of people on the train. In the current position, the train would be able to move up, left or right, but not down.

Each step of the train can be though of as a step of the leader plus a movement of the rest of the train to the previous position of the person immediately ahead of them on the train.

In this model, a new person can only join the train if they are in a valid position relative to the last person on the train. Valid positions are those adjacent to the last person on the train and that allow it to move according to the rules above. See the figure below for an example of valid positions to join the train (highlighted in green) at a given time.


<img src="images/train_add_person.png" alt="Human train definition" width="350"/>

### Classes and methods

Implement two classes to represent a person (`Person`) and a train (`Train`).

#### `Person` class (1 point)

In `Person`, you'll need to store the following information:
* `x`: the person's x coordinate;
* `y`: the person's y coordinate.

This class should have two methods, described below:

* Method **`is_at_position`** should:
  1. receive as arguments a pair of coordinates, `x` and `y`;
  2. check if the person's position matches that of the input;
  3. return a boolean.

* Method **`move_to_position`** should:
  1. receive as arguments a pair of coordinates, `x` and `y`;
  2. set the person's position to that of the input.

#### `Train` class

In `Train`, you'll need to store the following information:
* `people`: a list of the people on the train (which you may assume is always non-empty).

The `Train` class should have three methods, described below:

* Method **`is_at_position` (1 point)** should:
  1. receive as arguments a pair of coordinates, `x` and `y`;
  2. check if any of the people on the train has a position that matches that of the input;
  3. return a boolean.
  
  You may (and are strongly encouraged to) use the method `is_at_position` of each `Person` on the train.


* Method **`move` (2 points)** should:
  1. receive as argument a string identifying the direction to move &mdash; `"u"`, `"d"`, `"l"` or `"r"`, corresponding to a movement up, down, left or right, respectively;
  2. move the train in the desired direction, if possible; this operation should move each person on the train, *i.e.*, update their positions;
  3. return a boolean, which should be `True` if the movement was possible or `False` if it wasn't.
  
  You may (and are strongly enouraged to) use the method `move_to_position` of each `Person` on the train, as well as the method `is_at_position` of the `Train` class.


* Method **`add_person` (1 point)** should:
  1. receive as argument a `Person` object representing a person that wishes to join the train;
  2. add the person **to the end of the train**, if possible;
  3. return a boolean, which should be `True` if the input person could be added to the train or `False` if it couldn't.

<!-- grade_id: cell-25e4812b96a967d8, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Person:
    def __init__(self, x, y):
        self.x = x
        self.y = y
    
    def is_at_position(self, x, y):
        return self.x == x and self.y == y
    
    def move_to_position(self, x, y):
        self.x = x
        self.y = y
    
class Train:
    def __init__(self, people):
        self.people = people
    
    def is_at_position(self, x, y):
        return any([person.is_at_position(x, y) for person in self.people])        
        
    def move(self, direction):
        
        def follow_leader(people):
            for i in range(len(self.people) - 1, 0, -1):
                people[i].move_to_position(people[i-1].x, people[i-1].y)
        
        if direction == "r":
            # check that the train won't collide with itself
            if self.is_at_position(self.people[0].x + 1, self.people[0].y):
                return False
            
            # update all positions except the leader's
            follow_leader(self.people)
            
            # update leader's position to the right (x increases by 1)
            self.people[0].move_to_position(self.people[0].x + 1, self.people[0].y)
            
        elif direction == "l":
            if self.is_at_position(self.people[0].x - 1, self.people[0].y):
                return False
            follow_leader(self.people)
            self.people[0].move_to_position(self.people[0].x - 1, self.people[0].y)
        
        elif direction == "u":
            if self.is_at_position(self.people[0].x, self.people[0].y + 1):
                return False
            follow_leader(self.people)
            self.people[0].move_to_position(self.people[0].x, self.people[0].y + 1)

        elif direction == "d":
            if self.is_at_position(self.people[0].x, self.people[0].y - 1):
                return False
            follow_leader(self.people)
            self.people[0].move_to_position(self.people[0].x, self.people[0].y - 1)
        
        return True
    
    def add_person(self, new_person):
        
        # check that the new person's position does not collide with the train
        if self.is_at_position(new_person.x, new_person.y):
            return False
        
        last_person = self.people[-1]
        
        # check that the new person's position allows them to join the train
        if (
            new_person.is_at_position(last_person.x + 1, last_person.y) or
            new_person.is_at_position(last_person.x - 1, last_person.y) or
            new_person.is_at_position(last_person.x, last_person.y + 1) or
            new_person.is_at_position(last_person.x, last_person.y - 1)
        ):
            # add person to the train
            self.people.append(new_person)
        
        return True
    ### END SOLUTION
```

<!-- grade_id: cell-92c4e256e70cfb10, points: 5, grade: true, solution: false -->
```python
score = 0

try:
    person = Person(x=0, y=0)
    assert (person.x, person.y) == (0, 0), "The position of `person` is wrong."
    
    is_person_at_00 = person.is_at_position(0, 0)
    assert is_person_at_00 is True, "The verification of `person`'s position is wrong."
    
    is_person_at_10 = person.is_at_position(1, 0)
    assert is_person_at_10 is False, "The verification of `person`'s position is wrong."
    
    is_person_at_01 = person.is_at_position(0, 1)
    assert is_person_at_01 is False, "The verification of `person`'s position is wrong."
    
    person.move_to_position(x=2, y=-3)
    assert (person.x, person.y) == (2, -3), "The position of `person` after a movement is wrong."
except AssertionError as e:
    print(e)
    pass
else:
    score += 1

try:
    train = Train(people=[Person(0, 0), Person(1, 0)])
    
    assert isinstance(train.people, list), "The list of people in `train` has an incorrrect type."
    assert all([isinstance(person, Person) for person in train.people]), "The list of people in `train` should contain only instances of `Person`."
    assert len(train.people) == 2, "The list of people in `train` has an incorrect length."
    
    assert train.is_at_position(0, 0) is True, "The verification of `train`'s position is wrong."
    assert train.is_at_position(1, 0) is True, "The verification of `train`'s position is wrong."
    assert train.is_at_position(0, 1) is False, "The verification of `train`'s position is wrong."
except AssertionError as e:
    print(e)
    pass
else:
    score += 1

try:
    person_0, person_1 = Person(0, 0), Person(1, 0)
    train = Train(people=[person_0, person_1])
    
    move_return = train.move("u")
    assert move_return is True, "A valid movement of `train` has not been allowed."
    assert (train.people[0].x, train.people[0].y) == (0, 1), "The position of a person on the train is wrong."
    assert (train.people[1].x, train.people[1].y) == (0, 0), "The position of a person on the train is wrong."
    assert id(person_0) == id(train.people[0]), "`Person` objects should not be created / deleted when the train moves."
    assert id(person_1) == id(train.people[1]), "`Person` objects should not be created / deleted when the train moves."
    
    move_return = train.move("u")
    assert move_return is True, "A valid movement of `train` has not been allowed."
    assert (train.people[0].x, train.people[0].y) == (0, 2), "The position of a person on the train is wrong."
    assert (train.people[1].x, train.people[1].y) == (0, 1), "The position of a person on the train is wrong."
    assert id(person_0) == id(train.people[0]), "`Person` objects should not be created / deleted when the train moves."
    assert id(person_1) == id(train.people[1]), "`Person` objects should not be created / deleted when the train moves."
    
    move_return = train.move("l")
    assert move_return is True, "A valid movement of `train` has not been allowed."
    assert (train.people[0].x, train.people[0].y) == (-1, 2), "The position of a person on the train is wrong."
    assert (train.people[1].x, train.people[1].y) == (0, 2), "The position of a person on the train is wrong."
    assert id(person_0) == id(train.people[0]), "`Person` objects should not be created / deleted when the train moves."
    assert id(person_1) == id(train.people[1]), "`Person` objects should not be created / deleted when the train moves."
    
    move_return = train.move("r")
    assert move_return is False, "An invalid movement of `train` has been incorrectly allowed."
    assert (train.people[0].x, train.people[0].y) == (-1, 2), "The position of a person on the train is wrong."
    assert (train.people[1].x, train.people[1].y) == (0, 2), "The position of a person on the train is wrong."
    assert id(person_0) == id(train.people[0]), "`Person` objects should not be created / deleted when the train moves."
    assert id(person_1) == id(train.people[1]), "`Person` objects should not be created / deleted when the train moves."
    
    train = Train(people=[Person(0, 1), Person(0, 0), Person(1, 0), Person(1, 1), Person(1, 2)])
    move_return = train.move("r")
    assert move_return is False, "An invalid movement of `train` has been incorrectly allowed."
    
    
except AssertionError as e:
    print(e)
    pass
else:
    score += 2

try:
    person_0, person_1 = Person(-1, 2), Person(0, 2)
    train = Train(people=[person_0, person_1])
    
    person_2 = Person(0, 1)
    add_return = train.add_person(person_2)
    assert add_return is True, "A person in a valid position could not be added to `train`."
    assert len(train.people) == 3, "The length of `train` after a person was added to it is wrong."
    assert (train.people[0].x, train.people[0].y) == (-1, 2), "The position of a person on the train is wrong."
    assert (train.people[1].x, train.people[1].y) == (0, 2), "The position of a person on the train is wrong."
    assert (train.people[2].x, train.people[2].y) == (0, 1), "The position of the person added to the train is wrong."
    assert id(person_0) == id(train.people[0]), "Only one `Person` objects should not be created when a person joins the train."
    assert id(person_1) == id(train.people[1]), "Only one `Person` objects should not be created when a person joins the train."
    
    add_return = train.add_person(Person(0, 1))
    assert add_return is False, "An invalid addition of a person to `train` has been incorrectly allowed."
    assert len(train.people) == 3, "The length of `train` after an invalid addition of a person to it is wrong."
    
    add_return = train.add_person(Person(0, 2))
    assert add_return is False, "An invalid addition of a person to `train` has been incorrectly allowed."
    assert len(train.people) == 3, "The length of `train` after an invalid addition of a person to it is wrong."
    
except AssertionError as e:
    print(e)
    pass
else:
    score += 1

if score == 0:
    raise AssertionError("Not enough correct answers to score points.")

print(f"Your score is {score} / 5 in Exercise 3.")

score
```

Output:
```
Your score is 5 / 5 in Exercise 3.
5
```

### Exercise 4

Consider that you have a 3 x 3 matrix that represents a completed game of [Tic-tac-toe](https://en.wikipedia.org/wiki/Tic-tac-toe), where players O and X attempt to place three marks in a horizontal, vertical or diagonal row.

Each element of this matrix should be labeled with `"circle"` or `"cross"`, depending on whether they were played by players O and X, respectively, or `"none"` if the element has not been chosen by either player. An example of such matrix is shown below.

<code>
[["circle", "none", "cross"],
 ["none", "circle", "circle"],
 ["cross", "cross", "cross"]]
</code>

#### Part I (2 points)

Consider that the matrix described above is corrupted, such that each element may contain labels with mixed upper and lower case characters and spaces and/or underscores (`"_"`) before and/or after the corresponding string.

For example, the matrix shown above could have been corrupted to become

<code>
[["CIRcle", "_none", "cross"],
 ["NONE", " cIrcle", "circle"],
 ["crOSS", "cRoss_", " cross"]]
</code>

Write a function, named `clean_ttt_matrix` that receives as argument a numpy matrix that represents a **corrupted Tic-tac-toe** game, and returns a clean version of that game.

<!-- grade_id: cell-5a7eaacbb7097d35, points: 0, grade: false, solution: true -->
```python
def clean_ttt_matrix(corrupted_matrix):
    """
    Cleans a matrix representing a corrupted Tic-tac-toe game.

    Parameters:
    corrupted_matrix (np.ndarray): Corrupted matrix.

    Returns:
    uncorrupted_matrix (np.ndarray): Uncorrupted version of the input corrupted matrix.
    """
    
    ### BEGIN SOLUTION
    uncorrupted_matrix = np.copy(corrupted_matrix)
    
    for i in range(3):
        for j in range(3):
            uncorrupted_matrix[i, j] = uncorrupted_matrix[i, j].lower().strip().replace("_", "")

    return uncorrupted_matrix
    ### END SOLUTION
```

<!-- grade_id: cell-2eb102f666a51257, points: 2, grade: true, solution: false -->
```python
corrupted_matrix = np.array([
    ["CIRcle", "_none", "cross"],
    ["NONE", " cIrcle", "circle"],
    ["crOSS", "cRoss_", " cross"]
])

uncorrupted_matrix = clean_ttt_matrix(corrupted_matrix)
assert isinstance(uncorrupted_matrix, np.ndarray)
assert uncorrupted_matrix.shape == (3, 3)
assert uncorrupted_matrix[0, 0] == "circle"
assert uncorrupted_matrix[0, 1] == "none"
assert uncorrupted_matrix[1, 0] == "none"
assert uncorrupted_matrix[2, 1] == "cross"
assert uncorrupted_matrix[2, 2] == "cross"

print('All asserts passed.')
```

Output:
```
All asserts passed.
```

#### Part II (2 point)

Implement a function called `determine_ttt_winner` that, given a clean matrix representing a Tic-tac-toe game, determines who won the game.

The function should return `"O"` or `"X"` if player O or X won the game, respectively, or `None` if there are no winners.

**Hint:** the numpy functions `np.all`, `np.diagonal` and `np.fliplr` may be useful in this exercise. 

<!-- grade_id: cell-d2b1996422298e46, points: 0, grade: false, solution: true -->
```python
def determine_ttt_winner(ttt_game):
    """
    Determines the winner of a completed Tic-tac-toe game.

    Parameters:
    ttt_game (np.array): Matrix representing the game.

    Returns:
    winner (str or None): The winner of the game ("O" or "X") or `None` if there are no winners.
    """
    
    ### BEGIN SOLUTION
    # Check if rows are all filled with the same string
    for i in range(3):
        if np.all(ttt_game[i, :] == "circle"):
            return "O"
        elif np.all(ttt_game[i, :] == "cross"):
            return "X"
    
    # Do the same check for columns
    for j in range(3):
        if np.all(ttt_game[:, j] == "circle"):
            return "O"
        elif np.all(ttt_game[:, j] == "cross"):
            return "X"
    
    # Do the same check for the diagonal
    if np.all(np.diagonal(ttt_game) == "circle"):
        return "O"
    elif np.all(np.diagonal(ttt_game) == "cross"):
        return "X"
    
    # Do the same check for the anti-diagonal
    if np.all(np.diagonal(np.fliplr(ttt_game)) == "circle"):
        return "O"
    elif np.all(np.diagonal(np.fliplr(ttt_game)) == "cross"):
        return "X"
        
    return None
    ### END SOLUTION
```

<!-- grade_id: cell-8da0084a934a316d, points: 2, grade: true, solution: false -->
```python
ttt_game_0 = np.array([
    ["circle", "none", "cross"],
    ["none", "circle", "circle"],
    ["cross", "cross", "cross"]
])

winner_0 = determine_ttt_winner(ttt_game_0)
assert winner_0 == "X"

ttt_game_1 = np.array([
    ["none", "none", "circle"],
    ["cross", "cross", "circle"],
    ["circle", "cross", "circle"]
])

winner_1 = determine_ttt_winner(ttt_game_1)
assert winner_1 == "O"

ttt_game_2 = np.array([
    ["circle", "cross", "cross"],
    ["cross", "circle", "circle"],
    ["cross", "circle", "circle"]
])

winner_2 = determine_ttt_winner(ttt_game_2)
assert winner_2 == "O"

ttt_game_3 = np.array([
    ["circle", "none", "cross"],
    ["none", "cross", "circle"],
    ["cross", "circle", "none"]
])

winner_3 = determine_ttt_winner(ttt_game_3)
assert winner_3 == "X"

ttt_game_4 = np.array([
    ["cross", "circle", "circle"],
    ["circle", "cross", "cross"],
    ["cross", "circle", "circle"]
])

winner_4 = determine_ttt_winner(ttt_game_4)
assert winner_4 is None

print('All asserts passed.')
```

Output:
```
All asserts passed.
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)