# SLU14 - Modules and Packages 

#### Start by importing the following module

```python
#used for evaluation
import utils
```

### Exercise 1: What is the difference between a class and a module?

a) There is no difference; 

b) A module is a set of files with the `.py` extension and a class is just one of those files;   

c) A module is a file with the `.py` extension which can contain classes or one class, but can also contains functions and variables;   

d) A class is a file with the `.ipynb` extension and a module with the `.py` extension;

<!-- grade_id: cell-ec299e14ecce1f27, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60bb3a2487965c64, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q1, "Wrong answer."
```

### Exercise 2: How can we run the module as a script?

a) Just create your module with the desired attributes only: classes, functions and/or variables;      

b) With no special organization, but with all the code inside `if __init__:`;   

c) With no special organization at all, because it automatically already runs as a script;   

d) We only have to make sure that the entry point (the part of the code where the program starts executing) should be inside  `if __name__=='__main__':`;

<!-- grade_id: cell-2f1a5dfccc2eee9b, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18368bb83c4a41b4, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q2, "Wrong answer."
```

### Exercise 3: If you create an empty new module (.py file) and print __name__. What will happen and why? 

a) It will give an error, because the module can not be empty;   

b) It will print ```__main__ ```, because python always sets the variable 'name' to 'main' before running the file.;

c) It will print ```__name__ ```, because that is exactly what you asked python to do.;

<!-- grade_id: cell-7a2c49000d940397, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cd2427e89b7707cb, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q3, "Wrong answer."
```

### Exercise 4: So, now you made your first module and import it in your second module. Why are you happy that you used `if __name__=="__main__"` in your first module?

a) It looks interesting in your first module to use it, but doesn't do anything for your second module;  

b) It is not possible to import the first module without the statement `if __name__=='__main__':`;   

c) Python checked, during the import, if the name is "main". This is true (as seen in the previous question) and therefor it will only use the attributes within the if-statement and not everything in the first module file;  

d)Python checked, during the import, if the name is "main". This is not true, since the name will be set as the name of the imported module. But because of the statement `if __name__=='__main__':`, python knows to set the name back to "main". Now it is clear that the second module is your main file;   

<!-- grade_id: cell-a4658c8feaa4f078, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a98124ab22debe2, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q4, "Wrong answer."
```

### Exercise 5: Inside the directory stores, create a module named bookstore with the following attributes:


- A `class` named `Book` with three `attributes`, `title`, `author`, and `price`. The title attribute should be a string representing the title of the book, the author attribute should be a string representing the author of the book, and the price attribute should be a float representing the price of the book. The class should also have the following method: `get_book_info()`: returns a string that contains the book's title, author, and price in the following format: "Title: [title], Author: [author], Price: [price]."
- A function named `get_total_price` that accepts a list of Book objects as an input and returns the total price of all the books in the list.
- A variable named `description` with the following string value: `"This is a module named bookstore.`"
- Now let's see if your module works and follow the steps below to see the total price of your new books. 

To evaluate your module, create a variable named done with the value True.

<!-- grade_id: cell-aea5631b25c21caa, points: 0, grade: false, solution: true -->
```python

book_list = [
    ("To Kill a Mockingbird", "Harper Lee", 2),
    ("The Great Gatsby", "F. Scott Fitzgerald", 5),
    ("Pride and Prejudice", "Jane Austen", 7),
    ("The Catcher in the Rye", "J.D. Salinger", 6),
    ("1984", "George Orwell", 15)
    ]
    

# Set done = True

# import your module

# use the book_list
 
# Calculate the total price of the books using the get_total_price function

# total_price =

# Print the total price

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-421c00edf62020d8, points: 7, grade: true, solution: false -->
```python
if done:
    try:
        from stores import bookstore
        store = True
    except ImportError:
        store = False
        
    if store:
        assert bookstore, "Did you store a module named bookstore inside stores folder?"
        this_bookstore = bookstore.Book("Hamlet", "William Shakespear", 10)
        assert this_bookstore.title == "Hamlet", "Did you set title as an attribute of class Book?"
        assert this_bookstore.author == "William Shakespear", "Did you set auteur as an attribute of class Book?"
        assert this_bookstore.get_book_info()== "Title: Hamlet, Author: William Shakespear, Price: 10", "Did you set get_book_info correctly as a method of class Book?"
        assert bookstore.get_total_price(book_list), "Did you create a function named get_total price?"
        assert bookstore.description == "this is a module named bookstore", "Did you set description well?"
        assert total_price == 35, "The total price should be 35 euro's"
        print("Great job! You are awesome :-).")
    else:
        print("In the directory 'stores' is no module called 'bookstore'. Try again or ask for help")

else:
    print("You have to create a variable, called 'done', with the value 'True'")
        
```

### Exercise 6: We love a good python course, right?! Import the class Course, from module course, and create a new object from that.

The `library` folder is in the current directory. Explore its contents.
Import the class Course , that belongs to module course. Do this import in a way that the Course class name has no prefix. This doesn't mean that we cannot have dots in our import statement.
Afterwards, create a new course, named this_course, with name=prep-course, author=ldssa and year=2023. 

<!-- grade_id: cell-8dafeb32a322444f, points: 0, grade: false, solution: true -->
```python
# from module1.module2.etc... import Class,etc..
#this_course = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac799d1b55a4cba, points: 4, grade: true, solution: false -->
```python
assert Course, "Course class was not imported with success"
assert isinstance(this_course, Course), "this_course is not of type Course"
assert this_course.name == "prep-course", "this_course has the wrong name"
assert this_course.author == "ldssa", "this_course has the wrong author"
assert this_course.year == 2023, "this_course has the wrong year"
```

### Exercise 7: What is a package in Python?

a) A single Python file (with a .py extension) is considered a package;   

b) A set of modules that share the same directory and this directory must include an `__init__.py` file;

c) The thing the postman delivers when you order something online;

d)  A set of modules that share the same directory and this directory should not include an `__init__.py` file;

<!-- grade_id: cell-17731565ade4c8d8, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed761f5a351a7094, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q8, "Wrong answer."
```

### Exercise 8: Exercise

Build a python package called `calculator` with the following directory structure:

``` python
calculator
|____main.py
|____
| |____addition.py
| |____subtraction.py

```

`addition.py` should define the following function:

```python
def add(a, b):
    return a + b
```

`subtraction.py` should define the following function:

```python
def subtract(a, b):
    return a - b
```

`main.py` should define the following statement:

```python

print(add(2, 3))
print(subtract(5, 1))
```

You should then be able to run the program in a terminal and get the output:

``` python
> python calculator/main.py
5
4

```

Beware:
* the directory diagram above may be missing something... 
* the main file may be missing something... 

Were you able to complete this exercise? (if not, ask for help!)

a) Yes   
b) No

<!-- grade_id: cell-8df2491f6b0042a0, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-839cdbbdbfebce43, points: 2, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q9, "Wrong answer."
```

### Exercise 9: Why are packages important to use for any Python user?

a) Packages in Python were created to make programming more difficult and confusing for beginners.;

b) It is not really necessary anymore, nor is it important, but it is just the way it has always been done.;

c) Packages in Python are important because they allow you to hide your code from other developers and prevent them from using it.;

d) They are used to organize and distribute modules and we love organized. A lot of time you import modules that contains other modules. And this way it is easier to manage and organize your code. ;

<!-- grade_id: cell-471e5327dcc8e7b7, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d78f1dc248bb974, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q10, "Wrong answer."
```

### Exercise 10:  What is the unix command used to manage packages in a python environment?

a) cat   
b) pip   
c) package   
d) lib   

<!-- grade_id: cell-6d4d775081c8be95, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-342a8848d6c5e6fe, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q11, "Wrong answer."
```

# Submit your work!

To submit your work, [follow the steps here, in the step "Grading the Exercise Notebook"!](https://github.com/LDSSA/ds-prep-course-2023#22---working-on-the-learning-units)