# SLU02 - Exercise notebook

## Start by importing this package

```python
# Just for evaluating the results.
import math
import hashlib
```

## Exercise 1: Data Types

## 1.1) Scientific notation

<!-- grade_id: cell-a5bd49d98abf2410, points: 0, grade: false, solution: true -->
```python
# Assign the speed of light (2.99*10^8 m/s) to the variable speed_light.
# Use the scientific notation and 2 decimal places.

# speed_light = ...
### BEGIN SOLUTION

# Here is the actual solution
speed_light = 2.99E8

### END SOLUTION

# Calculate the distance (in meters) that a light particle travels in 15 seconds.
# Assign the distance travelled by the particle (in meters) to the variable distance_travelled.

# distance_travelled = 
### BEGIN SOLUTION

# Here is the actual solution
distance_travelled = speed_light * 15

### END SOLUTION
```

<!-- grade_id: cell-6b465d1d500666a0, points: 2, grade: true, solution: false -->
```python
assert isinstance(speed_light, float), "Review the data type of speed_light."
assert math.isclose(speed_light,2.99E8,abs_tol=1E5), "Check your speed!"
assert isinstance(distance_travelled, float), "Review the data type of distance_travelled."
assert math.isclose(distance_travelled,4.485E9,abs_tol=1E5), "Did you use the correct time in the calculation?"
print('Nicely done!')
```

Output:
```
Nicely done!
```

## 1.2) Data type conversion

The data types of several variables got mixed up. Convert them to the target type.

**Do not change** the values of the original variables.

### 1.2.1) To string

<!-- grade_id: cell-cf0b9cd761180962, points: 0, grade: false, solution: true -->
```python
# Assign the value of the variable size converted into a string to the variable size_string.
size = 1.78

# size_string = ...
### BEGIN SOLUTION

# Here is the actual solution
size_string = str(size)

# Hash something that checks some subset of the answer.
# In this example, it is checking the first element.
# This can help check subsets of the student answer in order
# to give more granular feedback.
print(
    'The size string:',
    hashlib.sha256(
        bytes(size_string, encoding='utf8')
    ).hexdigest()
)

### END SOLUTION
```

Output:
```
The size string: 5d48c7c6dce082d397fecf99b8f1ac7fa7adf02bd38fcbb1f7f9f13fba7a2d45
```

<!-- grade_id: cell-ac16bba863ed00af, points: 2, grade: true, solution: false -->
```python
size_string_hash = '5d48c7c6dce082d397fecf99b8f1ac7fa7adf02bd38fcbb1f7f9f13fba7a2d45'
assert isinstance(size_string, str), "Review the data type of size_string."
assert size_string_hash == hashlib.sha256(bytes(size_string, encoding='utf8')).hexdigest(), "The value of size_string is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 1.2.2)  To integer

<!-- grade_id: cell-ca00a15a43742c1c, points: 0, grade: false, solution: true -->
```python
# Assign to variable age_int the value of variable age converted into an integer. 
age = "38"

# age_int = ...
### BEGIN SOLUTION

# Here is the actual solution
age_int = int(age)

### END SOLUTION
```

<!-- grade_id: cell-43f5849625865cc0, points: 2, grade: true, solution: false -->
```python
assert isinstance(age_int, int), "Review the data type of age_int."
assert math.isclose(age_int, 38, abs_tol=0.1), "The value of age_int is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 1.2.3)  To float

<!-- grade_id: cell-cce49b7314add616, points: 0, grade: false, solution: true -->
```python
# Assign to the variable weight_float the value of the variable weight converted into a float. 
weight = "77.4"

# weight_float = ...
### BEGIN SOLUTION

# Here is the actual solution
weight_float = float(weight)

### END SOLUTION
```

<!-- grade_id: cell-97305f11f776d1a5, points: 2, grade: true, solution: false -->
```python
assert isinstance(weight_float, float), "Review the data type of weight_float."
assert math.isclose(weight_float, 77.4, abs_tol=0.01), "The value of weight_float is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

# Exercise 2: Operations

## 2.1) Arithmetic operations

Solve the math problems and assign the results to the variables.

Use arithmetic operations in your solutions.

If the answer has no decimal part, assign it as an **integer**. Otherwise leave it as a **float**.

### 2.1.1)  Basic arithmetics

<!-- grade_id: cell-68e58a91a765c6cb, points: 0, grade: false, solution: true -->
```python
# To study Data Science, Jane read 2 books. 
# She read the first one in one week, reading 27 pages everyday. 
# She read the second book in 11 days, reading 24 pages everyday. 
# What is the total number of pages that Jane read? 

# pages_read = ...
### BEGIN SOLUTION

# Here is the actual solution
pages_read = 7 * 27 + 11 * 24

### END SOLUTION
```

<!-- grade_id: cell-b09b6a3df43346e0, points: 2, grade: true, solution: false -->
```python
assert isinstance(pages_read, int), "Review the data type of pages_read."
assert math.isclose(pages_read, 453, abs_tol=0.001), "The value of pages_read is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 2.1.2) More basic arithmetics

<!-- grade_id: cell-0f95f2283fa7b5ed, points: 0, grade: false, solution: true -->
```python
# Steve used his savings of 2500 euros to do some purchases.
# He spent 3/4 of his savings on furniture. 
# He then spent 1/2 of his remaining savings on a fridge. 
# How much money had Steve left from his savings after the purchases?

#savings_left =
### BEGIN SOLUTION

# Here is the actual solution
savings_left = (2500 * 1 / 4) / 2

### END SOLUTION
```

<!-- grade_id: cell-b15473e30437404d, points: 2, grade: true, solution: false -->
```python
assert isinstance(savings_left, float), "Review the data type of savings_left."
assert math.isclose(savings_left, 312.5, abs_tol=0.2), "The value of savings_left is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 2.1.3) Even more arithmetics

<!-- grade_id: cell-4a8b61ca7fd9597b, points: 0, grade: false, solution: true -->
```python
# John, Mary, Louis and Silvia have 177 cards to use in a game. 
# They want to share them equally. 
# How many cards should each one take and how many cards will be left in the deck?

# cards_per_player = ...
# cards_left = 
### BEGIN SOLUTION

# Here is the actual solution
cards_per_player = 177 // 4
cards_left = 177 % 4

### END SOLUTION
```

<!-- grade_id: cell-80de177e3edd27bc, points: 2, grade: true, solution: false -->
```python
assert isinstance(cards_per_player, int), "Review the data type of cards_per_player."
assert math.isclose(cards_per_player, 44, abs_tol=0.001), "The value of cards_per_player is incorrect."
assert isinstance(cards_left, int), "Review the data type of cards_left."
assert math.isclose(cards_left, 1, abs_tol=0.001), "The value of cards_left is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 2.1.4) Shortcut arithmetic operators

<!-- grade_id: cell-ca40d76b6b0fd084, points: 0, grade: false, solution: true -->
```python
# Louis has invested 2700 euros in a savings account. 
# The investment gives returns of 5% of the value invested every year.
# Louis reinvests all the returns into the same savings account.
# How much money is there in the savings account at the end of the 3rd year?


# TIP: you can use as many auxiliary variables as you need. 
# Or change the value of a variable after the initial assignment.

# final_savings =
### BEGIN SOLUTION

# Here is the actual solution
final_savings = 2700
final_savings *= 1.05
final_savings *= 1.05
final_savings *= 1.05

### END SOLUTION
```

<!-- grade_id: cell-822b5ec7dfe02776, points: 2, grade: true, solution: false -->
```python
assert isinstance(final_savings, float), "Review the data type of final_savings."
assert math.isclose(final_savings, 3125.5875, abs_tol=0.1), "The value of final_savings is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

## 2.2) String operations

### 2.2.1) Concatenation

<!-- grade_id: cell-0a458560de5f3c4e, points: 0, grade: false, solution: true -->
```python
# Concatenate the joke variable with 3 times the response variable, leaving a space between each string.
# The result should be: "How does Darth Vader like his toast? On the dark side. Ha Ha Ha"
joke = "How does Darth Vader like his toast? On the dark side."
response = "Ha"
# joke_reception = ...

### BEGIN SOLUTION

# Here is the actual solution
joke_reception = joke + 3 * (" " + response)

# Hash something that checks some subset of the answer.
# In this example, it is checking the first element.
# This can help check subsets of the student answer in order
# to give more granular feedback.
print(
    'The size string:',
    hashlib.sha256(
        bytes(joke_reception, encoding='utf8')
    ).hexdigest()
)

### END SOLUTION
```

Output:
```
The size string: 3b544b8a3e3c38076931e145008372476700160b6c106dea0aede521192a5a88
```

<!-- grade_id: cell-65f0dc204b4b0812, points: 2, grade: true, solution: false -->
```python
joke_reception_hash = '3b544b8a3e3c38076931e145008372476700160b6c106dea0aede521192a5a88'
assert isinstance(joke_reception, str), "Review the data type of joke_reception."
assert joke_reception_hash == hashlib.sha256(bytes(joke_reception, encoding='utf8')).hexdigest(), "The value of joke_reception is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

### 2.2.2) String formatting

<!-- grade_id: cell-d2b9abfbcaf35a41, points: 0, grade: false, solution: true -->
```python
# Use one of the string formatting methods to insert the values of the variables into a string.
# Write the string in a single statement but DO NOT write the quantities "by hand" and DO NOT use concatenation.
# The result should be: "In the fridge there are 12 yogurts, 0.5 liters of milk and 8 eggs."
yogurt = 12
milk = 0.5
eggs = 8

# inventory = "In the fridge (...)"

### BEGIN SOLUTION

# Here is the actual solution
inventory = f"In the fridge there are {yogurt} yogurts, {milk} liters of milk and {eggs} eggs."

# Hash something that checks some subset of the answer.
# In this example, it is checking the first element.
# This can help check subsets of the student answer in order
# to give more granular feedback.
print(
    'The size string:',
    hashlib.sha256(
        bytes(inventory, encoding='utf8')
    ).hexdigest()
)

### END SOLUTION
```

Output:
```
The size string: d48f0e70729026e365046f03606e62d2dfab6b3b1102300e38efbeae63ce156d
```

<!-- grade_id: cell-9f9722689d8a954e, points: 2, grade: true, solution: false -->
```python
inventory_hash = 'd48f0e70729026e365046f03606e62d2dfab6b3b1102300e38efbeae63ce156d'
assert isinstance(inventory, str), "Review the data type of inventory."
assert inventory_hash == hashlib.sha256(bytes(inventory, encoding='utf8')).hexdigest(), "The value of inventory is incorrect."
print('Nicely done!')
```

Output:
```
Nicely done!
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)