# SLU01 Command Line - Guide to the Exercise Notebook

```python
# For evaluation purposes, please run this cell to import this package:
import hashlib
```

# Exercise 0: Getting help

In this dummy exercise, we will explain how to use the Exercise notebook. Each exercise follows the same structure:
- A markdown cell with the instructions for the problem - here it is explained what is the task to be performed and the expected result. Some necessary information might be displayed in the code cell as well.
- A code cell - in this cell you are asked to write the solution for the exercise.
- A cell with multiple `assert` statements (test cell) - this cell is used to test your solution.

You can consider this cell to be the instructions cell of exercise 0.

For each exercise you are expected to read the instructions, plan your solution, write the solution into the code cell, execute the code cell, and then execute the cell with the `assert` statements, **in this order.** Then:

- If everything is correct, the solution executes without any issues, it is accepted by the test cell and no error message appears. ✔️

- If you are shown an error message when **executing your code (on the code cell)**, it means that you made some **programming mistake(s)** that you should fix. ❌

- If you are shown an error message when **executing the test cell**, it means that even if the **code is correctly executed** by Python, it **does not produce the results that were expected**. ❌

Each `assert` tests a specific component of the solution. If a tests fails, an **AssertionError** message appears below the test cell that may provide some clues about what the issue might be. 

Below is an example where the second test failed, and the `assert` provides the hint "Check your speed!". The failed assert is indicated with an arrow. To solve this issue, you should read the hint and find out why your code is not producing the expected results. **After** making changes to the code, you execute the code cell again and **then** execute the test cell. You repeat this until you get the right solution.

<img src="media/assert_example.JPG" >

In the code cell, the instructors might provide some guidance or create variables or functions to be used in the exercise. You **should not change or delete** this. If you do so by mistake, you might need to copy the original text from the repository.

There are specific places for you to write your solutions. These places start with the comment "YOUR CODE HERE" and are followed by a `raise NotImplementedError()` statement as shown here:

```python
# YOUR CODE HERE
raise NotImplementedError()
```

You have to **delete the `raise NotImplementedError()` statement** and write your solution after the comment. The solution may take as many lines as needed. You always write one statement per line.

```python
# YOUR CODE HERE
a = 20
b = 22
solution = a + b
```

Be careful when writing your solution to not overwrite the instructor's text.

For this dummy exercise, the answers are provided. These answers are, however, **wrong**. They will produce programming errors or the tests will fail.

The objective is for you to read the error message, review the code and fix the error(s).

If you can't fix the issue or don't understand the error message you should:
- Read the error message and see if it's an obvious mistake.
- Review the code line by line to see if you can identify the error.
- Follow the *How to get help* section of SLU02 (this section is about resolving programming errors, you will not need it here in SLU01)
- Use a search engine to search for the error.
- Then, if you put in significant effort but are still stuck, you can ask the instructors for help on Slack.

## 0.1) Fix this solution

<!-- grade_id: cell-f74fc8192524b04c, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the integer value of 1731 to the variable radius.


### BEGIN SOLUTION
radius = 1731
### END SOLUTION
#This is where you put your solution. It is empty for the other exercises.
#I have attempted to solve the exercise but I made some mistakes. Silly me.
#Can you fix the following statement and perform the requested task?

radis=1131
```

<!-- grade_id: cell-3195fd0ef8dd6f52, points: 0, grade: true, solution: false -->
```python
radius_hash = '05758cd3875ad2171484c0026ccbb8adc210cd2d852407e3c7af1b751f35fdd6'
assert isinstance(radius,int), "The variable radius is not an integer. Re-read the task carefully."
assert radius_hash == hashlib.sha256(bytes(str(radius), encoding='utf8')).hexdigest(), "The value of radius is incorrect. Re-read the task carefully."
```

## 0.2) Now fix this solution

<!-- grade_id: cell-518166189407dfd9, points: 0, grade: false, solution: true -->
```python
# Assign the string "Beetlejuice!Beetlejuice!Beetlejuice!" to the variable ghost.

### BEGIN SOLUTION
ghost = 3 * "Beetlejuice!"

#Just for the passing the CircleCI test. The student won't have variable beetlejuice assigned.
beetlejuice = 1337 
### END SOLUTION
# Again, try to fix my solution.

ghosts = 3 + "beetlejuice"


# Additional text that the instructor might include. DO NOT DELETE ME!
```

Output:
```
TypeError                                 Traceback (most recent call last)
Input In [4], in <cell line: 11>()
      7 beetlejuice = 1337 
      8 ### END SOLUTION
      9 # Again, try to fix my solution.
---> 11 ghosts = 3 + "beetlejuice"
     14 #Additional text that the instructor might include. DO NOT DELETE ME!

TypeError: unsupported operand type(s) for +: 'int' and 'str'
```

<!-- grade_id: cell-bffeb92aae7d2585, points: 0, grade: true, solution: false -->
```python
ghost_hash = '5c2da3f9fb9ef74b5fe884c41977af0b1f29a5a0209d77cb2f313869ab8ab8fa'
assert 'ghost' in globals(), "There is no variable ghost. Review the code."
assert isinstance(ghost, str), "ghost is not a string. Review the code."
assert ghost_hash == hashlib.sha256(bytes(ghost, encoding='utf8')).hexdigest(), "ghost is a string but has the wrong value. Which function can you use to show values on the screen?"
```

Now that you know how to go about the Exercise notebook, it's time to tackle the exercises. Good luck! 🍀