# SLU01 Command Line - Exercise Notebook

```python
# For evaluation purposes, please run this cell to import this package:
import hashlib
```

All the exercises in this notebook are multiple choice questions. The choices are numbered by letters from A to D. In your answer, indicate the letter of the correct option. Write your answers in this way, with the letter enclosed in apostrophes (example for answer to exercise 1):

```Python
exercise_1='A'
```

### Exercise 1: Where is home?

<img src="media/tortoise_home.png" width="400"/>

Imagine that you have a Linux user called `tortoise`.What is the correct path to this user's home directory?

&emsp;A: `/root`   
&emsp;B: `/tortoise`   
&emsp;C: `/root/tortoise`   
&emsp;D: `/home/tortoise`

<!-- grade_id: cell-a894a58ceb39eaa2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_1 ='D'
### END SOLUTION

#exercise_1 =
```

<!-- grade_id: cell-5a2440598289b1e3, points: 2, grade: true, solution: false -->
```python
ex_1 = '50c9e8d5fc98727b4bbc93cf5d64a68db647f04f'
assert isinstance(exercise_1, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_1 == hashlib.sha1(bytes(exercise_1, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 2: Going home
Now for the same `tortoise` user, if the current working directory is `/var/lib/systemd/`, how do we go to this user's home directory?

&emsp;A: `pwd`   
&emsp;B: `cd /home`   
&emsp;C: `cd ~`   
&emsp;D: `cd tortoise`

<!-- grade_id: cell-29db176b41ae1c07, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_2 ='C'
### END SOLUTION

#exercise_2 =
```

<!-- grade_id: cell-e4ea53657d321c19, points: 2, grade: true, solution: false -->
```python
ex_2 = '484d52691fcadbfabec5a318d1cf9692c7f293cbc8c1d5f22b2d839b'
assert isinstance(exercise_2,str), "The answer should be enclosed in quotes or apostrophes."
assert ex_2 == hashlib.sha224(bytes(exercise_2, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 3: Going to places
Now for the same `tortoise` user, if the current working directory is `/var/lib/systemd/` and we run `cd ~/../..`, what will be the new current working directory? 

&emsp;A: `/home/tortoise`   
&emsp;B: `/home/tortoise/home/tortoise`   
&emsp;C: `~`   
&emsp;D: `/`

<!-- grade_id: cell-29db176b41ae1c00, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_3 ='D'
### END SOLUTION

#exercise_3 =
```

<!-- grade_id: cell-e4ea55657d321c19, points: 2, grade: true, solution: false -->
```python
ex_3 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(exercise_3, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_3 == hashlib.sha224(bytes(exercise_3, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 4: New name

How do we rename a non empty folder from `tortoise` to `tortoise-new`?

&emsp;A: `rename tortoise tortoise-new`   
&emsp;B: `rn tortoise tortoise-new`   
&emsp;C: `mv tortoise tortoise-new`   
&emsp;D: `mv -r tortoise tortoise-new`

<!-- grade_id: cell-befccac89f6c99df, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_4 ='C'
### END SOLUTION

#exercise_4 =
```

<!-- grade_id: cell-8e25f4730f887ad4, points: 2, grade: true, solution: false -->
```python
ex_4 = '484d52691fcadbfabec5a318d1cf9692c7f293cbc8c1d5f22b2d839b'
assert isinstance(exercise_4, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_4 == hashlib.sha224(bytes(exercise_4, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 5: Be fruitful and multiply
`tortoise` is not happy. This user has two directories in the current working directory, `baby-tortoise` and `tortoise-family`. They need to copy the `baby-tortoise` directory into the `tortoise-family` directory. 

Which command should they use?

&emsp;A: `mv baby-tortoise tortoise-family`   
&emsp;B: `cp baby-tortoise tortoise-family`   
&emsp;C: `cp -a baby-tortoise tortoise-family`   
&emsp;D: `cp -r baby-tortoise tortoise-family`

<!-- grade_id: cell-d629c9437fa027f7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_5 ='D'
### END SOLUTION

#exercise_5 =
```

<!-- grade_id: cell-9b757ee86b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_5 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(exercise_5, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_5 == hashlib.sha224(bytes(exercise_5, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 6: Build a house, plant a tree

<img src="media/directories.png" width="400"/>

Create the file system tree in the image above, with the indicated directories and files.

Were you able to do it?

&emsp;A: YES   
&emsp;B: NO, and I'll ask some questions on Slack to help me figure this out!

<!-- grade_id: cell-d629c9437fa026f7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_6 ='A'
### END SOLUTION

#exercise_6 =
```

<!-- grade_id: cell-9b757ef86b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_6 = '5cfe2cddbb9940fb4d8505e25ea77e763a0077693dbb01b1a6aa94f2'
assert isinstance(exercise_6, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_6 == hashlib.sha224(bytes(exercise_6, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 7: Remove all traces

<img src="media/tortoise.jpg" width="400"/>

Before you start questioning yourself whether you have a weird fixation over tortoises, you might want to remove some directories and files from the current working directory.

What is the correct command to remove the `tortoisephilia` folder (which has some tortoise pictures)?

&emsp;A: `drop tortoisephilia`   
&emsp;B: `rm tortoisephilia/*`   
&emsp;C: `remove tortoisephilia`   
&emsp;D: `rm -r tortoisephilia`

<!-- grade_id: cell-d629c9437fa126f7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_7 ='D'
### END SOLUTION

#exercise_7 =
```

<!-- grade_id: cell-9b757ef81b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_7 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(exercise_7, str), "The answer should be enclosed in quotes or apostrophes."
assert ex_7 == hashlib.sha224(bytes(exercise_7, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 8: Sneak peak
Our `tortoise` user created a `tortoise_codes.py` file with some Python code (as all good tortoises do), now they need to check what is stored in the file in a simple and easy way.

What is the correct command for that procedure?

&emsp;A: `cat tortoise_codes.py`   
&emsp;B: `read tortoise_codes.py`   
&emsp;C: `ls tortoise_codes.py`   
&emsp;D: `touch tortoise_codes.py`

<!-- grade_id: cell-1792bbe6ebc8564d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_8 ='A'
### END SOLUTION

#exercise_8 =
```

<!-- grade_id: cell-d57ae086b71ffed4, points: 2, grade: true, solution: false -->
```python
ex_8 = 'ad14aaf25020bef2fd4e3eb5ec0c50272cdfd66074b0ed037c9a11254321aac0729985374beeaa5b80a504d048be1864'
assert isinstance(exercise_8,str), "The answer should be enclosed in quotes or apostrophes."
assert ex_8 == hashlib.sha384(bytes(exercise_8, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 9: Emil and the detectives

This is the output of a certain bash command. From the following options, which one could it be?

```bash
-rw-rw-r-- 1 maria maria 10451 Apr  1 22:49 'Exercise notebook guide.ipynb'
-rw-rw-r-- 1 maria maria 18943 Apr  2 00:24 'Exercise notebook.ipynb'
-rw-rw-r-- 1 maria maria 25338 Apr  1 22:26 'Learning notebook.ipynb'
drwxrwxr-x 2 maria maria  4096 Apr  1 21:00  media
-rw-rw-r-- 1 maria maria  1252 Apr  1 21:27  README.md
-rw-rw-r-- 1 maria maria  3338 Mär 21 20:03  submit.py
```

&emsp;A: `head ~/projects/ds-prep-course-2023/Week\ 1/SLU01\ - Command\ Line`   
&emsp;B: `cat ~/projects/ds-prep-course-2023/Week\ 1/SLU01\ - Command\ Line`   
&emsp;C: `ls -l`   
&emsp;D: `ls -tl`

<!-- grade_id: cell-82043c053fb1c39d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_9 ='C'
### END SOLUTION

#exercise_9 =
```

<!-- grade_id: cell-fa61be6aa4097dad, points: 2, grade: true, solution: false -->
```python
ex_9 = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
assert isinstance(exercise_9,str), "The variable should be a string."
assert ex_9 == hashlib.sha256(bytes(exercise_9, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

### Exercise 10: Count 'em all

The `~/projects/ds-prep-course-2023/Week 01/SLU01 - Command Line` directory on your computer contains these files and directories:

```bash
Exercise notebook guide.ipynb
Exercise notebook.ipynb
Learning notebook.ipynb
media
README.md
submit.py
```

How many of them are text files? Use a `bash` command to find out.

&emsp;A: 1   
&emsp;B: 2   
&emsp;C: 3   
&emsp;D: 5

<!-- grade_id: cell-8204dc053fb1c39d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_10 ='D'
### END SOLUTION

#exercise_10 =
```

<!-- grade_id: cell-fa61be6aa4d97dad, points: 2, grade: true, solution: false -->
```python
ex_10 = '3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43'
assert isinstance(exercise_10,str), "The answer should be enclosed in quotes or apostrophes."
assert ex_10 == hashlib.sha256(bytes(exercise_10, encoding='utf8')).hexdigest(), "Not the correct option."
print("That's right, way to go!")
```

Output:
```
That's right, way to go!
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)